# IFN619 - Assessment 1 - Part A

Answer the questions below to demonstrate your understand of foundational techniques. Each question relates to a technique that has been addressed in the first 4 modules.

Each question will be marked separately, so if you are unsure about a question, simply try to answer it with your best possible attempt. You can run the *test* cells (after each answer cell) to check for simple errors.

<span style="color:red;">__DO NOT ADD OR REMOVE ANY CELL BLOCKS. PROVIDE YOUR ANSWERS IN THE INDICATED SPACES ONLY, AND DO NOT DELETE THE ### LINES.__</span>

## 1. Read a CSV file

After importing the required dataframes library as `pd`, read the contents of file `mountain_heights.csv` into a dataframe called `mountain_heights`. Use the mountain code for the index of the dataframe. Using a dataframe property, output the dataframe's rows and columns in the format `(rows,cols)`. 

**This solution requires 3 lines of code.**

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q1-tests, points: 1, grade: true, solution: false -->
```python
"Check that the dataframe has been loaded and rows and columns output"
assert(type(mountain_heights)==type(pd.DataFrame()))
assert(_==(10,4))

```

## 2. Fix missing data

Using a single line of code, assign the string `Antarctica` to the appropriate cell of the dataframe by using the index and column labels.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q2-tests, points: 1, grade: true, solution: false -->
```python
"Check that there are no NaN cells"
assert(mountain_heights['Country'].isnull().values.any()==False)

```

## 3. Finding significant data

Find the mountains that are higher than the average height of all mountains in the data. Start by getting the mean and assigning it to a variable called `mean_height`. Using the format `df[c]>x` as the conditional for the rows, create a new dataframe called `highest_mountains` which have heights greater than the mean. Without using `print()`, display the value of mean_height in the output of the cell.

**This solution requires 3 lines of code**


<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q3-tests, points: 1, grade: true, solution: false -->
```python
"Check highest_mountains shape, and output of cell"
assert(highest_mountains.shape==(3,4))
assert(_==6060.9)

```

## 4. Visualise

First, sort the `mountain_heights` dataframe by height in descending order and assign the result to a dataframe called `mountains_sorted`. Use a single line of code.

Then in a second single line of code, create a bar chart of the `mountains_sorted` dataframe. Assign the result of the plot to a variable called `mplot`, and include the following information:

* Title: Mountain Heights
* X Axis label: Mountain Code
* Y Axis label: Height (m)


<!-- grade_id: q4-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q4-tests, points: 1, grade: true, solution: false -->
```python
"Check mountains_sorted, and that a chart has been created"
assert(list(mountains_sorted.index) == ['EVR', 'ACG', 'DNL', 'LGN', 'KLM', 'CRC', 'ELB', 'ORZ', 'VNM', 'JAY'])
assert(type(mplot).__name__=='AxesSubplot')

```

## 5. Read a text file

Using the format `with open('myfile','r') as file:`, read the contents of file `system.log` into a single string and assign it to a variable called `syslog`. Display the contents of the variable in the output of the cell.

<!-- grade_id: q5-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q5-tests, points: 1, grade: true, solution: false -->
```python
"Check reading using file"
import io
assert(type(file)==type(io.TextIOWrapper(io.BytesIO())))
"Check read as string"
assert(type(syslog)==type(""))

```

## 6. Structure text

Using the `syslog` string, create a list with each line being an element of the list. Ensure that lines do not include any newline characters. Assign the list to a variable called `sysloglist`. Iterate over this list and for each entry, create a list of date, time, code, and message. For example: `['2021-03-01','17:00:00','INFO','System startup complete']`. Add each created list to a new list called `syslogtable`. After iterating over `sysloglist`, the `syslogtable` should be a list of lists. Display this in the output of the cell.

<!-- grade_id: q6-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q6-tests, points: 1, grade: true, solution: false -->
```python
"check list created without newlines"
assert(len(sysloglist)==9)
assert('\n' not in sysloglist)
"check table displayed"
table = _
assert(len(table)==9)
assert(len(table[8])==4)

```

## 7. Convert to JSON

Create a list called `headings` with the strings `date`, `time`, `code`, and `message` corresponding to the values in the `syslogtable`. Get the 7th entry of syslog table and turn it into a dictionary with the headings as keys. Assign it to a variable called `syslogentry`. Hint: use the zip function `dict(zip(keys,values))`. Convert the dictionary to a JSON string using the `json` library. Call the result `jsonentry` and display in the output of the cell.


<!-- grade_id: q7-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q7-tests, points: 1, grade: true, solution: false -->
```python
"Check headings, syslogentry and jsonentry"
assert(len(headings)==4)
assert(type(syslogentry)==type({}))
assert(type(jsonentry)==type(""))
```

## 8. Read a JSON file

Using the `url` below and the `requests` library, get data from the server and read it into a dictionary called `jupyterdata`. Display in the output of this cell a string with the format `version: value` where `value` is the value in the dictionary corresponding to the key `version`.

<!-- grade_id: q8-answer, points: 0, grade: false, solution: true -->
```python
url = "https://jupyter-oncampus.qutanalytics.io/hub/api"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q8-tests, points: 1, grade: true, solution: false -->
```python
"Check jupyterdata"
assert(type(jupyterdata)==type({}))
```

## 9. Create a function

Create a function called `get_json()` that when passed a URL, returns a dictionary created from the resulting JSON retrieved from that URL. For the purposes of this task, the function should work the url from question 8. Run the function to output the version of this server.

<!-- grade_id: q9-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q9-tests, points: 1, grade: true, solution: false -->
```python
"Check function returns value from url"
assert(get_json(url)['version']=='1.3.0')
```

## 10. Save JSON to file

Create a modified version of your `get_json()` function to save the resulting JSON data to a file called `output.json`. Call the function `save_json()` (Note: It should accept the URL as before).

<!-- grade_id: q10-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q10-tests, points: 1, grade: true, solution: false -->
```python
"Check function, and that file is created after running"
assert(callable(save_json))
import os.path
save_json(url)
assert(os.path.isfile('output.json'))
with open('output.json','r') as checkfile:
    check = json.loads(checkfile.read())
assert(check=={'version': '1.3.0'})
```